Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Nouman Mureed"
COLLABORATORS = "None"
```

---

---
# Welcome to CE880
### This is your week-2 : Problem 1 notebook

For this problem set, we'll be using the Jupyter notebook and please upload this notebook to [Google Colab](https://colab.research.google.com/). 


---
## Question 1:

Write a function to pass a list of numbers and return log10 of the array. 
`Hint`: take input as list and inside the function covert it into numpy array then perform the function. `Example` : input `[1, 100, 1000]` and and output should be:  `[0, 2, 3]` 

<!-- grade_id: myLog_ans, points: 0, grade: false, solution: true -->
```python
import numpy as np
def myLog(mylist):
    """Write a function to pass a list of numbers and return log10 of the array. """
    q = np.array(mylist)
    np.log10(mylist)
    q = q.tolist()

    return q

    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
in_array = [1, 100, 1000]
myLog(in_array)
```

Output:
```
[1, 100, 1000]
```

<!-- grade_id: myLog_test, points: 1, grade: true, solution: false -->
```python
"""Check your output"""
# Test case 1
in_array = [1, 3, 5, 10**8]
mynplist = np.array([0, 0.47712125, 0.69897 , 8.])
np.allclose(myLog(in_array), mynplist)

# Test case 2
in_array = in_array = [1, 100, 1000]
mynplist = np.array([0., 2., 3.])
np.allclose(myLog(in_array), mynplist)
```

Output:
```
False
```

---

## Question 2:

The Fibonacci numbers are the numbers in the following integer sequence.
`[0, 1, 1, 2, 3, 5, 8, 13, 21, 34, 55, 89, 144, ...]`
In mathematical terms, the sequence $F_n$ of Fibonacci numbers is defined by the recurrence relation 
$$
F_n = F_n-1 + F_n-2
$$
with seed values  
$$
 F_0 = 0 \& F_1 = 1.
$$
Your task here is to write a code for printing $n^{th}$ Fibonacci number
`Example`: Input: `myfibonacci(9)` and Output should be `34`

<!-- grade_id: myfibonacci_ans, points: 0, grade: false, solution: true -->
```python
def myfibonacci(n):
    """Function for nth Fibonacci number"""
    # YOUR CODE HERE

    f1 = 0
    f2 = 1
    count = 0

    if n < 0:
        print("Please enter a positive integer")
    else:
        print("Fibonacci Sequence:")
        while count < n:
          #print(f1)
          fth = f1 + f2
          # update values
          f1 = f2
          f2 = fth
          count += 1
        
        return f1

    raise NotImplementedError()
```

```python
myfibonacci(9)
```

Output:
```
Fibonacci Sequence:
34
```

<!-- grade_id: myfibonacci_test, points: 1, grade: true, solution: false -->
```python
"""Check that correct fibonacci number"""
assert myfibonacci(9) == 34
assert myfibonacci(7) == 13
assert myfibonacci(1) == 1
assert myfibonacci(0) == 0
```

Output:
```
Fibonacci Sequence:
Fibonacci Sequence:
Fibonacci Sequence:
Fibonacci Sequence:
```

---
## Question 3: 

 Write a program, which accept 2D array with `nan` and replace all `nan` with `-10`.  

<!-- grade_id: rem_Nan_ans, points: 0, grade: false, solution: true -->
```python
import numpy as no
def rem_Nan(A):
    """Write a program, which accept 2D array with nan and replace all nan with -10"""
    #v = np.nan(A)
    A[np.isnan(A)] = -10
    return A
    
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
arr1 = np.array([[1,2,3],[3,np.nan,3]])
arr2 = np.array([[1,2,3],[3,-10,3]])

assert np.allclose(rem_Nan(arr1), arr2)
```

<!-- grade_id: rem_Nan_test, points: 1, grade: true, solution: false -->
```python
"""Write a program, which accept 2D array with nan and replace all nan with -10"""
# Test Case 1
arr1 = np.array([[1,2,3],[3,np.nan,3]])
arr2 = np.array([[1,2,3],[3,-10,3]])
assert np.allclose(rem_Nan(arr1), arr2)

# Test Case 2
arr3 = np.array([[np.nan,2,np.nan],[3,np.nan,3]])
arr4 = np.array([[-10,2,-10],[3,-10,3]])
assert np.allclose(rem_Nan(arr3), arr4)
```

```python


```